```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

Output:
```
d:\miniforge\envs\dsc80\lib\site-packages\otter\export\__init__.py:6: UserWarning: pkg_resources is deprecated as an API. See https://setuptools.pypa.io/en/latest/pkg_resources.html. The pkg_resources package is slated for removal as early as 2025-11-30. Refrain from using this package or pin to Setuptools<81.
  import pkg_resources
```

# Lab 2 – DataFrames and Grouping

## DSC 80, Spring 2026

### Due Date: Wednesday, April 15th at 11:59PM

## Instructions

Welcome to the second DSC 80 lab this quarter!

Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Unlike in DSC 10, labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc80-2024-fa).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.** More details on its usage are given at the bottom of this notebook.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import os
import io
from pathlib import Path
import pandas as pd
import numpy as np
```

<div class="alert alert-block alert-danger">

<b>The only question in this lab that you're allowed to use a loop in is Question 3.</b> There, you may use a <code>for</code>-loop to loop over the columns in the input DataFrame, but not the rows. <b>If you use a <code>for</code>-loop or <code>while</code>-loop in any other question, you may lose points!</b>

</div>

## Part 1: Tricky Pandas 🤔

Sometimes, `pandas` gives you weird outputs that you may not expect. The next question walks you through a few examples that might surprise you. 

### Question 1

The following subparts all require you to define a function and return a number that is the answer to a multiple-choice question. You may need to write code and experiment with DataFrames to arrive at your answers.

#### `trick_me`

`trick_me` should not take any arguments. 
<br>

Inside the function:

* Create a DataFrame named `tricky_1` that has three columns labeled `'Name'`, `'Name'`, and `'Age'`. `tricky_1` should have 5 rows; the values are up to you.
* Save the DataFrame to a `.csv` file called `'tricky_1.csv'` without the index.
* Now create another DataFrame, named `tricky_2`, by reading in the file `'tricky_1.csv'`. What are your observations?

  1. It was not possible to create a DataFrame with the duplicate columns.
  2. `tricky_1` and `tricky_2` have the same column names.
  3. `tricky_1` and `tricky_2` have different column names.
   
Your function should return `1`, `2`, or `3`, answering the above question.

<font color='red'>**Hints:** Dictionaries can only contain unique keys, so you <b>cannot</b> use them to declare a DataFrame with duplicate column names. Find other ways to create a DataFrame with duplicate column names.</font>

<br>
  
#### `trick_bool`
`trick_bool` should not take any arguments.

To determine the correct answer from the list below, you should follow the steps outlined by experimenting in **the notebook** (or in the Terminal by running `python`). Outside the function:

* Create a DataFrame named `bools` that has four columns: `True`, `True`, `False`, `False`. Each column name should be Boolean.
* `bools` should have 4 rows; the values are up to you.
* Predict the shape of the DataFrame that results by running each of the three lines of code below. Pick a corresponding answer from the given list. Your function should return a list with three numbers, one for each line.
* You should be able to answer without running any code, but feel free to run code to check your answer.
* **Your function should not do anything other than return a hardcoded list.**

```py
bools[True]
bools[[True, True, False, False]]
bools[[True, False]]
```
    
Answer choices:
1. DataFrame: 2 columns, 1 row
2. DataFrame: 2 columns, 2 rows
3. DataFrame: 2 columns, 3 rows
4. DataFrame: 2 columns, 4 rows
5. DataFrame: 3 columns, 1 rows
6. DataFrame: 3 columns, 2 rows
7. DataFrame: 3 columns, 3 rows
8. DataFrame: 3 columns, 4 rows
9. DataFrame: 4 columns, 1 rows
10. DataFrame: 4 columns, 2 rows
11. DataFrame: 4 columns, 3 rows
12. DataFrame: 4 columns, 4 rows
13. Error

**Hints:** Refer to the previous hint for `trick_me`.

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
trick_ans = trick_bool()
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Part 2: Summary Statistics 📊

In this question you will define two general purpose functions that make it easy to qualitatively assess the contents of a DataFrame.

### Question 2

Complete the implementation of the function `population_stats`, which takes in a DataFrame `df` and returns a DataFrame indexed by the columns of `df`, with the following columns:
* `'num_nonnull'`, which contains the number of non-null entries in each column.
* `'prop_nonnull'`, which contains the proportion of entries in each column that are non-null.
* `'num_distinct'`, which contains the number of distinct non-null entries in each column.
* `'prop_distinct'`, which contains the proportion of non-null entries that are distinct in each column.
       
For example, if `df` has a column named `'ages'` with the following elements (note that `np.nan` is a null value):
       
```py
[2, 2, 2, np.nan, 5, 7, 5, 10, 11, np.nan]
```

Then:
- `'num_nonnull'` is 8, and `'prop_nonnull'` is $\frac{8}{10}$ = 0.8.
- There are six distinct entries, `[2, 5, 7, 10, 11, np.nan]`, but only 5 of them are non-null. So the number of distinct non-null entries, `'num_distinct'`, is 5.
- There are 5 distinct non-null entries, and there are 8 total non-null entries, so `'prop_distinct'` is $\frac{5}{8}$ = 0.625.

Putting it all together, `population_stats(df).loc['ages']` should be a Series containing the numbers 8, 0.8, 5, and 0.625.

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
pop_data = np.random.choice(range(10), size=(100, 4))
df_pop = pd.DataFrame(pop_data, columns='A B C D'.split())
out_pop = population_stats(df_pop)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### Question 3
    
Complete the implementation of the function `most_common`, which takes in a DataFrame `df` and a number `N` and returns a DataFrame of the `N` most-common values and their counts for each column of `df`. Any column with fewer than `N` distinct values should contain `np.nan` in those entries.

For example, consider the DataFrame shown on the left. This DataFrame is a subset of `salaries`, a larger DataFrame containing information on employees in the City of San Diego. The subset below contains two of the original columns: `'Job Title'` which contains job titles for employees, and `'status'` which denotes whether the employee works a full time position (`'FT'`) or a part time position (`'PT'`). On the right, the return value of `most_common(salaries, N=5)` is shown.

You can assume that there are no ties in our hidden tests.

<table><tr>
    <td><img src="data/imgs/dataframe.png" width="90%"/></td>
    <td><img src="data/imgs/most_common.png" width="90%"/></td>
</tr></table>

***Note***: Remember, to access values in a Series based on their integer position, including when slicing for the first `N` values in a Series, you **must** use `.iloc` followed by square brackets. If you just use square brackets and don't use `.iloc`, you may not see the results you expect!

***Hint***: You may find that initializing an empty DataFrame with `N` rows and adding columns to it is useful in your implementation.

<div class="alert alert-block alert-danger">

<b>Remember, the only question in this lab that you're allowed to use a loop in is Question 3 – that's this question.</b> Here, you may use a <code>for</code>-loop to loop over the columns in the input DataFrame (<code>df</code>), but not the rows. <b>If you use a <code>for</code>-loop or <code>while</code>-loop in any other question, you may lose points!</b>

</div>

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
common_data = np.random.choice(range(10), size=(100, 2))
common_df = pd.DataFrame(common_data, columns='A B'.split())
common_out = most_common(common_df, N=3)
common_out
```

Output:
```
A_values  A_counts  B_values  B_counts
0         2        15         6        14
1         1        13         2        14
2         8        12         4        12
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Part 3: Superheroes 🦸

The questions below analyze a dataset of superheroes found in the `data` directory. One of the datasets lists the attributes of each superhero, while the other is a *Boolean* DataFrame describing which superheroes have which superpowers. Note, the datasets contain information on both **good** superheroes, as well as **bad** superheroes (AKA villains).

If you took DSC 10 in Fall 2022, this dataset may seem familiar – it was used for the Final Project that quarter!

### Question 4

Let's start working with the `powers` dataset, which you can see in `data/superheroes_powers.csv`. 

Complete the implementation of the function `super_hero_powers`, which takes in a DataFrame like `powers` and returns a list with the following three entries:

1. The name of the superhero with the greatest number of superpowers.
2. Identify the most common superpower among superheroes who can fly, other than `'Flight'` itself..
3. The name of the most common superpower among superheroes with only one superpower.

You should **not** be hard-coding your answers in this question; your function should work on any DataFrame similar to `powers`. In each case, you can assume the answer is unique.

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
super_fp = Path('data') / 'superheroes_powers.csv'
powers = pd.read_csv(super_fp)
super_out = super_hero_powers(powers)
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

### Question 5

In the notebook, load in the dataset in `data/superheroes.csv` as a DataFrame and explore it. Call your `population_stats` function from Question 2 on the DataFrame. You should notice that there are very few actually null (`np.nan`) values, but there are many entries that **should** be null, because they're missing.

Complete the implementation of the function `clean_heroes`, which takes in a DataFrame like the one created from `superheroes.csv` and returns a new DataFrame with all of the missing values replaced with `np.nan`.

After cleaning the superheroes dataset with `clean_heroes`, run `population_stats` on it again. As a result of the cleaning, population_stats should show that there are more null values.

***Note***: Most of the work in this question is identifying how the missing values are stored in the DataFrame. The implementation of the function should only take one line.

```python
# Cell 1: 加载看看
heroes = pd.read_csv(Path('data') / 'superheroes.csv', index_col=0)
heroes.head(20)
```

Output:
```
name  Gender Eye color               Race Hair color  Height  \
0              A-Bomb    Male    yellow              Human    No Hair   203.0   
1          Abe Sapien    Male      blue      Icthyo Sapien    No Hair   191.0   
2            Abin Sur    Male      blue            Ungaran    No Hair   185.0   
3         Abomination    Male     green  Human / Radiation    No Hair   203.0   
4             Abraxas    Male      blue      Cosmic Entity      Black   -99.0   
5       Absorbing Man    Male      blue              Human    No Hair   193.0   
6         Adam Monroe    Male      blue                  -      Blond   -99.0   
7        Adam Strange    Male      blue              Human      Blond   185.0   
8            Agent 13  Female      blue                  -      Blond   173.0   
9           Agent Bob    Male     brown              Human      Brown   178.0   
10         Agent Zero    Male         -                  -          -   191.0   
11         Air-Walker    Male      blue      
```

```python
population_stats(heroes)
```

Output:
```
num_nonnull  prop_nonnull  num_distinct  prop_distinct
name                734      1.000000           715       0.974114
Gender              734      1.000000             3       0.004087
Eye color           734      1.000000            23       0.031335
Race                734      1.000000            62       0.084469
Hair color          734      1.000000            30       0.040872
Height              734      1.000000            54       0.073569
Publisher           719      0.979564            24       0.033380
Skin color          734      1.000000            17       0.023161
Alignment           734      1.000000             4       0.005450
Weight              732      0.997275           135       0.184426
```

```python
# don't change this cell -- it is needed for the tests to work
superheroes_fp = Path('data') / 'superheroes.csv'
heroes = pd.read_csv(superheroes_fp, index_col=0)
clean_out = clean_heroes(heroes)
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

Below, we have displayed the first 10 rows of the cleaned DataFrame.

```python
clean_out.head(10)
```

Output:
```
name  Gender Eye color               Race Hair color  Height  \
0         A-Bomb    Male    yellow              Human    No Hair   203.0   
1     Abe Sapien    Male      blue      Icthyo Sapien    No Hair   191.0   
2       Abin Sur    Male      blue            Ungaran    No Hair   185.0   
3    Abomination    Male     green  Human / Radiation    No Hair   203.0   
4        Abraxas    Male      blue      Cosmic Entity      Black     NaN   
5  Absorbing Man    Male      blue              Human    No Hair   193.0   
6    Adam Monroe    Male      blue                NaN      Blond     NaN   
7   Adam Strange    Male      blue              Human      Blond   185.0   
8       Agent 13  Female      blue                NaN      Blond   173.0   
9      Agent Bob    Male     brown              Human      Brown   178.0   

           Publisher Skin color Alignment  Weight  
0      Marvel Comics        NaN      good   441.0  
1  Dark Horse Comics       blue      good    65.0  
2          DC Comic
```

### Question 6

Using the **cleaned** superhero data, we will now generate some insights.

Complete the implementation of the function `super_hero_stats`, which takes no arguments and returns a list of length 6 containing your answers to the questions below. **Your answers should be hard-coded in the function.**

0. What is the name of the tallest `'Mutant'` with `'No Hair'`?
1. Among the publishers who have more than 5 characters, which publisher has the highest proportion of human characters? If there is a tie, return the publisher whose name is first alphabetically. We define a character to be human if their `'Race'` is exactly the string `'Human'`; for instance, a `'Race'` of `'Human / Radiation'` is non-human for the purposes of this question.
2. Among the characters whose `'Height'`s we know, who is taller on average – `'good'` characters or `'bad'` characters?
3. Which publisher has a greater proportion of `'bad'` characters – `'Marvel Comics'` or `'DC Comics'`?
4. Which `'Publisher'` that isn't `'Marvel Comics'` or `'DC Comics'` has the most characters? Consider all characters whose `'Publisher'` we know – that is, don't drop rows because they have null values in other columns.
5. There is only one character that is **both** more one standard deviation above the mean in height and more than one standard deviation below the mean in weight. What is their name?

***Note***: When calculating your answers, do not drop rows based on null values.

***Note***: Although you'll be writing code to find the answers, you should not include your code in your `.py` file. Just return a hard-coded list with your answers to the 6 questions; all 6 elements in the list should be strings.

```python
heroes = pd.read_csv(Path('data') / 'superheroes.csv', index_col=0)
df = clean_heroes(heroes)

q0 = df[(df['Race']=='Mutant') & (df['Hair color']=='No Hair')]\
        .sort_values('Height', ascending=False).iloc[0]['name']

pub_counts = df['Publisher'].value_counts()
big_pubs = pub_counts[pub_counts > 5].index
sub = df[df['Publisher'].isin(big_pubs)]
human_prop = sub.groupby('Publisher').apply(
    lambda g: (g['Race']=='Human').sum() / len(g)
)
q1 = human_prop.sort_values(ascending=False).sort_index().sort_values(ascending=False).index[0]

q1 = human_prop.sort_index().sort_values(ascending=False, kind='stable').index[0]

q2 = df.groupby('Alignment')['Height'].mean().loc[['good','bad']].idxmax()

mdc = df[df['Publisher'].isin(['Marvel Comics','DC Comics'])]
q3 = mdc.groupby('Publisher').apply(lambda g: (g['Alignment']=='bad').mean()).idxmax()

q4 = df[~df['Publisher'].isin(['Marvel Comics','DC Comics']) & df['Publisher'].notna()]\
        ['Publisher'].value_counts().idxmax()

h_mean, h_std = df['Height'].mean(), df['Height'].std()
w_mean, w_std = df['Weight'].mean(), df['Weight'].std()
q5 = df[(df['Height'] > h_mean + h_std) & (df['Weight'] < w_mean - w_std)]['name'].iloc[0]

print([q0, q1, q2, q3, q4, q5])
```

Output:
```
['Onslaught', 'George Lucas', 'bad', 'Marvel Comics', 'NBC - Heroes', 'Groot']
C:\Users\Ethan\AppData\Local\Temp\ipykernel_35032\2164627276.py:10: DeprecationWarning: DataFrameGroupBy.apply operated on the grouping columns. This behavior is deprecated, and in a future version of pandas the grouping columns will be excluded from the operation. Either pass `include_groups=False` to exclude the groupings or explicitly select the grouping columns after groupby to silence this warning.
  human_prop = sub.groupby('Publisher').apply(
C:\Users\Ethan\AppData\Local\Temp\ipykernel_35032\2164627276.py:20: DeprecationWarning: DataFrameGroupBy.apply operated on the grouping columns. This behavior is deprecated, and in a future version of pandas the grouping columns will be excluded from the operation. Either pass `include_groups=False` to exclude the groupings or explicitly select the grouping columns after groupby to silence this warning.
  q3 = mdc.groupby('Publisher').apply(lambda g: (g['Alignmen
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
stats_out = super_hero_stats()
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

## Part 4: High Potential Individuals 📈

Last year, the United Kingdom 🇬🇧 announced a new ["High Potential Individual" visa](https://www.lexology.com/library/detail.aspx?g=41fa64ec-9272-468c-bdcb-8002745a754f), which allows graduates of universities ranked in the Top 50 globally to move to the UK without a job lined up. This visa has been a subject of much debate, in part due to how much rankings play a role. (Rest assured, UCSD is on the list!)

In this section, you will analyze a dataset of university rankings, collected from  [here](https://www.kaggle.com/datasets/mylesoneill/world-university-rankings?datasetId=) (though we have pre-processed and modified the original dataset for the purposes of this question). Our version of the dataset is stored in `data/universities_unified.csv`.

Columns:
* `'world_rank'`: world rank of the institution
* `'institution'`: name of the institution
* `'national_rank'`: rank within the nation, formatted as `'country, rank'`
* `'quality_of_education'`: rank by quality of education
* `'alumni_employment'`: rank by alumni employment
* `'quality_of_faculty'`: rank by quality of faculty
* `'publications'`: rank by publications
* `'influence'`: rank by influence
* `'citations'`: rank by number of citations
* `'broad_impact'`: rank by broad impact
* `'patents'`: rank by number of patents
* `'score'`: overall score of the institution, out of 100
* `'control'`: whether the university is public or private
* `'city'`: city in which the institution is located
* `'state'`: state in which the institution is located

### Question 7

There are (still) a few aspects of the dataset we need to clean before it's ready for analysis.

#### `clean_universities`

Complete the implementation of the function `clean_universities`, which takes in the raw rankings DataFrame and returns a cleaned DataFrame, cleaned according to the following information:

- Some `'institution'` names contain `'\n'` characters (e.g. `'University of California\nSan Diego'`). Replace all instances of `'\n'` with `', '` (a comma and a space) in the `'institution'` column.

- Change the data type of the `'broad_impact'` column to `int`.

* Split `'national_rank'` into two columns, `'nation'` and `'national_rank_cleaned'`, where:
    * `'nation'` is the country (or its dependency) indicated in the first part of `'national_rank'`. 
        * Note that there are **3** countries that appear under different names for different schools. For all 3 of these countries, you should pick **the name that is longer** and use that name for every occurrence of the country. One of the 3 countries is **`'Czech Republic'`**, which also appears as **`'Czechia'`** – since these refer to the same country and `'Czech Republic'` is longer, all instances of either name should be replaced with `'Czech Republic'`. You need to find the other 2 countries on your own. 
        * As is mentioned below, your function will only be tested on the DataFrame in `data/universities_unified.csv`, so you only need to change these 3 country names.
    * `'national_rank_cleaned'` is the integer in the latter part of `'national_rank'`. Make sure that the data type of this column is `int`. 
    * Don't include the original `'national_rank'` column in the output DataFrame.
* Create a Boolean column `'is_r1_public'`. This column should contain `True` if a university is public and classified as R1 and `False` otherwise. Treat `np.nan`s as False. **Note that in the raw DataFrame, a university is classified as R1 if and only if it has non-null values in all of the following columns: `'control'`, `'city'`, and `'state'`.**
    - Read [this page](https://en.wikipedia.org/wiki/List_of_research_universities_in_the_United_States) to learn more about what it means for a university to be classified as R1.
    
**The only dataset your function will be tested on is `data/universities_unified.csv`; you don't need to worry about other hidden test sets.** In addition, please return a *copy* of the original DataFrame; don't modify the original.

<br>

Now, we can do some basic exploration.

#### `university_info`

Complete the implementation of the function `university_info`, which takes in the **cleaned** DataFrame outputted by `clean_universities` and returns the following values in a list:
* Among `'state'`s with three or more `'institution'`s in the dataset, the `'state'` whose universities have the lowest mean `'score'`.
* The proportion of the `'institution'`s in the top 100 for which the `'quality of faculty'` ranking is also in the top 100.
* The number of `'state'`s where at least 50% of the `'institution'`s are private (i.e. have an `'is_r1_public'` of `False`).
* The lowest-ranked (worst) `'institution'` in the world, according to `'world_rank'`, that is the highest-ranked (best) university in its nation (i.e., it has a `'national_rank_cleaned'` of 1).

You can assume there are no ties.

```python
df = pd.read_csv(Path('data') / 'universities_unified.csv')
nations = df['national_rank'].str.split(', ', n=1, expand=True)[0]
print(sorted(nations.unique()))
```

Output:
```
['Argentina', 'Australia', 'Austria', 'Belgium', 'Brazil', 'Bulgaria', 'Canada', 'Chile', 'China', 'Colombia', 'Croatia', 'Cyprus', 'Czech Republic', 'Czechia', 'Denmark', 'Egypt', 'Estonia', 'Finland', 'France', 'Germany', 'Greece', 'Hong Kong', 'Hungary', 'Iceland', 'India', 'Iran', 'Ireland', 'Israel', 'Italy', 'Japan', 'Lebanon', 'Lithuania', 'Malaysia', 'Mexico', 'Netherlands', 'New Zealand', 'Norway', 'Poland', 'Portugal', 'Puerto Rico', 'Romania', 'Russia', 'Saudi Arabia', 'Serbia', 'Singapore', 'Slovak Republic', 'Slovenia', 'South Africa', 'South Korea', 'Spain', 'Sweden', 'Switzerland', 'Taiwan', 'Thailand', 'Turkey', 'UK', 'USA', 'Uganda', 'United Arab Emirates', 'United Kingdom', 'United States', 'Uruguay']
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
fp = Path('data') / 'universities_unified.csv'
df = pd.read_csv(fp)
cleaned = clean_universities(df)
info = university_info(cleaned)
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

## Congratulations! You're done Lab 2! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.

Once you've finished the lab, you should open the command line and run, in the directory for this lab:

```
python lab-validation.py
```

**This will run all of the `grader.check` cells that you see in this notebook, but only using the code in `lab.py` – that is, it doesn't look at any of the code in this notebook. If all of your `grader.check` cells pass in this notebook but not all of them pass in your command line with the above command, then you likely have code in your notebook that isn't in your `lab.py`!**

You can also use `lab-validation.py` to test individual questions. For instance,

```
python lab-validation.py q1 q4 q7
```

will run the `grader.check` cells for Questions 1, 4, and 7 – again, only using the code in `lab.py`. [This video](https://www.loom.com/share/0ea254b85b2745e59322b5e5a8692e91?sid=5acc92e6-0dfe-4555-9b6a-8115b6a52f99) how to use the script as well.

Once `python lab-validation.py` shows that you're passing all test cases, you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

<div class="alert alert-block alert-danger">

<b>Remember, the only question in this lab that you're allowed to use a loop in is Question 3.</b> There, you may use a <code>for</code>-loop to loop over the columns in the input DataFrame, but not the rows. <b>If you use a <code>for</code>-loop or <code>while</code>-loop in any other question, you may lose points!</b>

</div>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> ans =  trick_me()
>>> ans == 1 or ans == 2 or ans == 3
True
```
```
>>> isinstance(trick_ans, list)
True
```
```
>>> isinstance(trick_ans[1], int)
True
```


Tests for q2 (points: None)
```
>>> out_pop.index.tolist() == ['A', 'B', 'C', 'D']
True
```
```
>>> cols = ['num_nonnull', 'prop_nonnull', 'num_distinct', 'prop_distinct']
>>> out_pop.columns.tolist() == cols
True
```
```
>>> bool((out_pop['num_distinct'] <= 10).all())
True
```
```
>>> bool((out_pop['prop_nonnull'] == 1.0).all())
True
```


Tests for q3 (points: None)
```
>>> common_out.index.tolist() == [0, 1, 2]
True
```
```
>>> common_out.columns.tolist() == ['A_values', 'A_counts', 'B_values', 'B_counts']
True
```
```
>>> bool(common_out['A_values'].isin(range(10)).all())
True
```


Tests for q4 (points: None)
```
>>> isinstance(super_out, list)
True
```
```
>>> len(super_out) == 3
True
```
```
>>> all([isinstance(x, str) for x in super_out])
True
```


Tests for q5 (points: None)
```
>>> sum(clean_out['Skin color'].isnull()) > sum(heroes['Skin color'].isnull())
True
```
```
>>> sum(clean_out['Weight'].isnull()) > sum(heroes['Weight'].isnull())
True
```


Tests for q6 (points: None)
```
>>> isinstance(stats_out[0], str)
True
```
```
>>> isinstance(stats_out[1], str)
True
```
```
>>> stats_out[2] in ['good', 'bad']
True
```
```
>>> stats_out[3] in ['Marvel Comics', 'DC Comics']
True
```
```
>>> isinstance(stats_out[4], str)
True
```
```
>>> isinstance(stats_out[5], str)
True
```


Tests for q7 (points: None)
```
>>> cleaned.shape[0] == df.shape[0]
True
```
```
>>> cleaned['nation'].nunique() == 59
True
```
```
>>> len(info) == 4
True
```
```
>>> all([True if isinstance(x,y) else x.dtype==np.dtype(y) for x, y in zip(info, [str, float, int, str])])
True
```
```
>>> (info[1] >= 0) & (info[1] <= 1)
True
```
